<a href="https://colab.research.google.com/github/Chonthicha006/MTH361-Python-Coding/blob/main/Lab1_BasicPython_Student.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Basic Python สำหรับ Data Science 




## Python <a id="Python"></a>
* High level programming language
* Script programing language
* Dynamically typed language

### เวอร์ชั่นของ Python


```python
!python --version
```

Output:
```
Python 3.9.16
```

## Basic Data Types
Python เองก็มีตัวแปรพื้นฐานเช่นเดียวกับภาษาอื่นๆ อาทิเช่น จำนวนเต็ม (int) จำนวนจริง (float) บูเลียน (boolean) และสตริง (string)



### Integer and Real Number 


```python
x=10
```

```python
type(x)
```

Output:
```
int
```

```python
y=20
```

```python
z=x+y
z

a=2.0
print(type(a))
a+x
```

Output:
```
<class 'float'>
12.0
```

```python
id(z)
```

Output:
```
139866446240976
```

### Object oriented programming



`x = 10 `
เป็นการประกาศค่าเริ่มต้นให้ x โดย x จะกลายเป็นค่า object ประเภท int ที่มีค่า 10 โดยเราจะเห็นได้จาก

`<class 'int'>`

โดยแนวคิดของ object จะประกอบไปด้วย 2 ส่วนคือ
1. Property
2. Method (function)

อย่างเช่น class เป็น EV (รถยนต์ไฟฟ้า)
Property : สีรถ, เครื่องยนต์, ล้อรถ
Method: เดินหน้า, ถอยหลัง, เลี่ยว

แนวทางการใช้งาน


``` python
car1 = EV            # สร้าง object car1
car1.color = 'red'   # กำหนด property
car1.forward()       # สั่งให้เคลื่อนที่ไปข้างหน้า
                     # สังเกตว่ามีเครื่องหมายวงเล็บปิดท้าย method ส่วน property จะไม่มี
```




```python

```

```python

```

### Boolean
เป็นข้อมูลอีกประเภทหนึ่งที่ใช้งานในเชิงตรรกศาสตร์ คือ มีค่าสองค่า True และ False ซึ่งจะใช้งานในเชิงการตัดสินใจในการเขียนโปรแกรม

```python
t = True
f = False
```

```python
print(type(t))
print(type(f))
#ตรวจสอบว่า t=f หรือไม่
t==f
t!=f
```

Output:
```
<class 'bool'>
<class 'bool'>
True
```

### String 


```python
str1='Hello'
str2="world"
str1
str2 #เป็นqoute
```

Output:
```
'world'
```

```python
print(str1) #ออกมาเป็นอักขระ
str3=str1+' '+str2
str3
```

Output:
```
Hello
'Hello world'
```

```python
print(str1+' '+str2)
```

Output:
```
Hello world
```

```python
print('x=%d' %x)
print('y=%d' %y)
```

Output:
```
x=10
y=20
```

```python
print('x=%d,y=%d' %(x,y))
print('x=%.2f,y=%f' %(x,y))
```

Output:
```
x=10,y=20
x=10.00,y=20.000000
```

### Python is a dynamicall typed language
https://www.analyticsvidhya.com/blog/2021/04/an-overview-of-python-memory-management/



```python
id(z)
```

Output:
```
139866446240976
```

```python
id(x)
```

Output:
```
139866446240336
```

## Control Flow Statement



###if

ลักษณะโครงสร้างของการใช้ if เป็นดังนี้


```python
if some_condition:
    statement1
    statement2
```

ให้สังเกตว่าหลัง some\_condition จะตามด้วย colon \( : \)

```python
money=20
if money >=15:
  print('ไปซื้อโค้กกัน')
  print('จะเหลือเงิน %d บาท' %(money-15))
```

Output:
```
ไปซื้อโค้กกัน
จะเหลือเงิน 5 บาท
```

```python

```

### If\-else

กรณีที่มี 2 เส้นทางที่ต้องดำเนินการที่แตกต่างกันหลังจากการตรวจสอบเงื่อนไข

```python
if some_condition:
    statement1
    statement2
    
else:    
    statement3
    statement4
```



```python
money=20
if money >=15:
  print('ดื่มน้ำเปล่า')
else:
  print('ดื่มโค้ก')
```

Output:
```
ดื่มน้ำเปล่า
```

```python

```

### if\-elif

``` python
if some_condition:    
    statement1
    statement2

elif some_condition:    
    statement3
    statement4

else:    
    statement5
    statement6
```


```python

```

```python

```

## Loops
ในบางครั้งเราต้องการให้โปรแกรมคอมพิวเตอร์ทำงานในลักษณะที่ซ้ำๆ การเขียนด้วย for loop ก็จะช่วยให้เราทำงานสะดวกขึ้น

### For 
โครงสร้างการเขียน 

``` python
for variable in something:    
    statement1
    statement2
```

```python

```

```python

```

### While 

วิธีการการใช้งาน while loop เป็นดังนี้

```python
while some_condition:
    statement1
    statement2
```

```python

```

```python

```

## Function
เรามาทำความรู้การประกาศฟังก์ชัน ในเบื้องต้นก่อนเพื่อจะได้ทำแบบฝึกหัด
การประกาศฟังก์ชันมีรูปแบบดังนี้

``` python 

# This is formatted as code
def functionName(arguments):
    statement1
    statement2
    return  var1
```

```python

```

```python

```

return ของฟังก์ชันนั้นจะมีหรือไม่ก็ได้ ดังตัวอย่างต่อไปนี้ จะไม่มีการ return ค่าส่งกลับมา

```python

```

```python

```

## Exerciser 1 
จงเขียนโปรแกรมเพื่อตรวจสอบว่าประโยคทั้ง 2 ประโยคนี้เท่ากันหรือไม่ โดยไม่สนใจเรื่องของตัวอักษรตัวเล็กหรือตัวใหญ่


```python
!pip install nose
```

<!-- grade_id: 7c3d5d, points: 0, grade: false, solution: true -->
```python
def checkStringEqual(a, b): # modify function name and parameters
    """
    Check a and b are the same sentense without care about the Upper or Lower charactors
    """
    # YOUR CODE HERE
    raise NotImplementedError()
        
    return result 

```

<!-- grade_id: 70ede0, points: 1, grade: true, solution: false -->
```python
a = 'I Love Dog'
b = 'I love dog'
checkStringEqual(a,b)
c = 'I love cat'
checkStringEqual(a,c)

# [Modify the tests below for your own problem]
# Check that squares returns the correct output for several inputs:
from nose.tools import assert_equal
assert_equal(checkStringEqual(a,b), True)


```

## Exercise 2

จงเขียนโปรแกรม "for" loop เพื่อบวกเลข ที่เก็บไว้ใน a โดยเพิ่มขึ้นครั้งละ 1 ไปจนถึง b  โดย a < b อย่างเช่น <br>
a = 1 <br>
b = 4 <br>
totalSum = 1 + 2 + 3<br>

```python
def totalSum(a,b):
    total = 0
    # YOUR CODE HERE
    raise NotImplementedError()

    return total

```

```python
# [Modify the tests below for your own problem]
# Check that squares returns the correct output for several inputs:
from nose.tools import assert_equal
assert_equal(totalSum(1,100), 4950)
```

# Containers <a id="Containers"></a>
Python นั้นมีตัวแปรพิเศษที่สามารถใส่ตัวแปรอื่นๆ เข้าไปข้างในได้ (เรียกตัวแปรพิเศษเหล่านี้ว่า container) ที่หลากหลายและใช้งานสะดวกมาก เช่น  lists, dictionaries, sets, and tuples.




## List
เป็นตัวแปรทีสามารถเก็บตัวแปรชนิดอื่นๆ ได้ โดยไม่จำกัดชนิด และไม่จำกัดจำนวนและไม่จำเป็นต้องมีการประกาศไว้ล่วงหน้า 


``` python
lstA = [1, 'a', 'b', 3]
```



```python

```

```python

```

**ข้อควรระวัง**
เนื่องจาก Python เป็น Dynamicalled Typed Language อย่างที่ได้กล่าวไว้ เราจะทำสิ่งผิดพลาดต่อไปนี้ได้


``` python
lstB = lstA
```



```python

```

```python

```

### Loop with List
loop กับ list เป็นสิ่งที่คู่กัน เราสามารถนำเอาทั้งสองสิ่งที่นี้มาใช้งานได้อย่างมีประสิทธิภาพดังนี้

```python

```

```python

```

บางกรณีเราจำเป็นต้องอ้างถึงตำแหน่งของข้อมูล ต้องใช้คำสั่ง enumerate


``` python
A = ['x', 'y', 'z']
a = enumerate(A)
list(a)
```



```python

```

```python

```

### Slicing 
การทำ slicing (สไลด์ซิ่ง) บน list คือการเข้าถึงตัวแปรหลายๆ ตัวบน list พร้อมๆ กัน ซึ่งเป็นคุณสมบัติที่ใช้งานได้สะดวกมากๆ 

```python

```

```python

```

## Dictionary
Container สุดยอดอีกตัวของ Python คือ dictionary (เรียกสั้นๆ ว่า dict) โดยจะประกอบไปด้วย key และ value ประกอบกันเป็นคู่ดังนี้ 

```python 
{key1 : value1, key2 : value2} 
```
โดย key และ value จะเป็นตัวแปรประเภทใดก็ได้   

```python

```

```python

```

### Loops with Dict
เราสามารถวนลูปใน dict ได้ง่ายๆ เหมือน list 

```python

```

```python

```

## Set
set เป็น container ที่มีคุณสมบัติเหมือน set ในวิชาคณิตศาสตร์ นั่นคือสมาชิกใน set จะไม่มีการเรียงลำดับ และจะไม่ซ้ำกัน เช่น


```
st = {'cat', 'dog'}
```

```python

```

```python

```

### Loops with Set
การวนลูปทำได้เช่นเดียวกับ list ยกเว้นแต่เราจะไม่ทราบลำดับของสมาชิกใน set

เช่น

```
animals = {'cat', 'dog', 'fish'}
for idx, animal in enumerate(animals):
    print('#%d: %s' % (idx + 1, animal))
```

```python

```

```python

```

## Tuple
Tuple คือตัวแปรประเภท **sequential** (ลำดับ) เหมือนกับ list ซึ่งจะทำ slicing ได้  
การใช้งานประกาศดังนี้ 

```
t = (1, 2)
```


```python

```

```python

```

มีจุดแตกต่างที่สำคัญสองเรื่อง 
1. เรื่องแรกคือเมื่อมีการสร้าง tuple แล้ว เราจะไม่สามารถเปลี่ยนค่าข้างใน tuple ด้วยเครื่อง `=` ได้  


```python

```

```python

```

2. ข้อแตกต่างเรื่องที่สองคือ เราสามารถใช้ tuple เป็น key ใน dict หรือสมาชิกใน set ได้ (list ไม่สามารถใช้งานในรูปแบบเหล่านี้ได้) 

```python

```

```python

```

## Exercise 3
จากข้อมูลที่อยู่ใน Dictionary ดังนี้ 

``` python
d = {'person': 2, 'cat': 4, 'spider': 8, 'horse': 4, 'bird' : 2 }
```

จงเขียนโปรแกรมเพื่อแสดงค่าให้ส่งรายชือสัตว์ที่จำนวนขาตามกำหนดเช่น <br>
x = liveLeg(8) <br>
print(x) 

'spider' 


<!-- grade_id: 7a9523, points: 0, grade: false, solution: true -->
```python
d = {'person': 2, 'cat': 4, 'spider': 8, 'horse': 4, 'bird' : 2 }

def liveLegs(legs):
    live = [];
    # YOUR CODE HERE
    raise NotImplementedError() 
    return live

```

<!-- grade_id: 605bb2, points: 1, grade: true, solution: false -->
```python
# [Modify the tests below for your own problem]
# Check that squares returns the correct output for several inputs:
from nose.tools import assert_equal
assert_equal(liveLegs(2), ['person', 'bird'])
assert_equal(liveLegs(8), ['spider'])

```